<a href="https://colab.research.google.com/github/Harischandra359/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_count(json_file):
    """
    json_file(str) -> int

    This function reads a JSON file and checks if it contains a 'entry' key.
    If the 'entry' key is present, it returns the number of entries under it.
    If the file does not exist, it prints a message and returns 0.
    If there is an error decoding the JSON file, it prints an error message and returns 0.

    >>> type(allergy_count(ALLERGIES_FILE))
    int

    >>> allergy_count(ALLERGIES_FILE)
    4
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if 'entry' in data:
                return len(data['entry'])
            else:
                return 0
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return 0
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON: {str(e)}")
        return 0
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
def patient_count(json_file):
    """
    json_file(str) -> int

    This function reads a JSON file containing medical data and checks if it contains an 'entry' key.
    If the 'entry' key is present, it iterates through the entries and extracts patient information.
    It counts the number of unique patients and returns that count.
    If the file does not exist, it prints a message and returns 0.
    If there is an error decoding the JSON file, it prints an error message and returns 0.

    >>> type(patient_count(ALLERGIES_FILE))
    int

    >>> patient_count(ALLERGIES_FILE)
    2
    """

    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            patients = set()
            if 'entry' in data:
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource']:
                        patients.add(entry['resource']['patient']['display'])
                return len(patients)
            else:
                return 0
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return 0
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON: {str(e)}")
        return 0
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
def allergy_per_patient(json_file):
    """
    json_file(str) -> dict

    This function reads a JSON file containing medical data and checks if it contains an 'entry' key.
    If the 'entry' key is present, it iterates through the entries and extracts patient information.
    It counts the number of allergies for each patient and stores the count in a dictionary.
    The dictionary is then returned, where keys are patient names, and values are the corresponding allergy counts.
    If the file does not exist, it prints a message and returns an empty dictionary.
    If there is an error decoding the JSON file, it prints an error message and returns an empty dictionary.

    >>> type(allergy_per_patient(ALLERGIES_FILE))
    dict

    >>> allergy_per_patient(ALLERGIES_FILE)
    {'Paul Boal': 1, 'Jason Argonaut': 3}
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            allergy_count = {}
            if 'entry' in data:
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource']:
                        patient_name = entry['resource']['patient']['display']
                        if patient_name in allergy_count:
                            allergy_count[patient_name] += 1
                        else:
                            allergy_count[patient_name] = 1
                return allergy_count
            else:
                return {}
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return {}
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON: {str(e)}")
        return {}
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(json_file):
    """
    json_file (str) -> list

    This function reads a JSON file containing medical data and checks if it contains an 'entry' key.
    If the 'entry' key is present, it iterates through the entries and extracts patient names, allergies, and reactions.
    It constructs a list of lists, where each inner list contains this information.
    If no reaction is found for an entry, the default value "No Reaction" is used.
    The resulting list is then returned.
    If the file does not exist, it prints a message and returns an empty list.
    If there is an error decoding the JSON file, it prints an error message and returns an empty list.

    >>> allergy_list(ALLERGIES_FILE)
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
    ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
    ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
    ['Paul Boal', 'PENICILLIN G', 'Bruising']]
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            result = []
            if 'entry' in data:
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource']:
                        patient_name = entry['resource']['patient']['display']

                        if 'substance' in entry['resource']:
                            allergy = entry['resource']['substance']['text']

                        reaction = "No Reaction"  # Default value if no reaction is found
                        if 'reaction' in entry['resource'] and len(entry['resource']['reaction']) > 0:
                            reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']

                        result.append([patient_name, allergy, reaction])
                return result
            else:
                return []
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return []
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON: {str(e)}")
        return []
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(json_file, patient, substance):
    """
    (json_file(str),str, str) -> str or None

    This function first calls the 'allergy_list' function to obtain a list of allergies with patient names and reactions.
    It then iterates through the list to find the entry that matches the provided patient and substance.
    If a matching entry is found, the associated reaction is returned.
    If no matching entry is found, the function returns None.

    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
    'Hives'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS')
    'Itching'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY')
    'Anaphylaxis'
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN')
    None
    >>> allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')
    'Bruising'
    """
    allergy_data = allergy_list(json_file)
    for entry in allergy_data:
        if entry[0] == patient and entry[1] == substance:
            return entry[2]
    return None

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

